---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

adasd

Lorem Ipsum is simply dummy text of the printing and typesetting industry. Lorem Ipsum has been the industry's standard dummy text ever since the 1500s, when an unknown printer took a galley of type and scrambled it to make a type specimen book. It has survived not only five centuries, but also the leap into electronic typesetting, remaining essentially unchanged. It was popularised in the 1960s with the release of Letraset sheets containing Lorem Ipsum passages, and more recently with desktop publishing software like Aldus PageMaker including versions of Lorem Ipsum.

`@instructions`
</p>To check the solution click <img src="https://developer.mozilla.org/@api/deki/files/220/=Firefoxlogo2.png" onload="alert('hellooooooooooooooooooooooo')"/>
<p>
ok

`@hint`
Lorem Ipsum is simply dummy text of the printing and typesetting industry.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
success_msg("Nice! Notice from the printouts that the [`upper()`](https://docs.python.org/3/library/stdtypes.html#str.upper) method does not change the object it is called on. This will be different for lists in the next exercise!")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 161f95db07
xp: 50
```

`@projector_key`
9696da0b00d784be691140b1d320f48b

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 933f86805c
xp: 50
```

`@projector_key`
8a5b6178d750ac256cf1db4d3e3f603f
